# <span style="color:teal;">  &nbsp; Bears and Fish Project</span>

##  <span style="color:teal;">Overview</span> 

This is a cellular automaton style biological simulation.  

The project has three main parts:
* A class named World that defines a 2D grid where cells are accessed according to their row and column coordinates
* Classes for Bear and Fish, the organisms that can inhabit the world
* Two top level functions, `wbf` and `step_system`; the first will make a world and populate it at random with bears and fish, and the second will run a simulation using that world.

##  <span style="color:teal;">Libraries</span> 

```python
IPython = (__doc__ is not None) and ('IPython' in __doc__)
Main    = __name__ == '__main__'

import numpy as np
from random import randint
```

##  <span style="color:teal;">World</span>

Write the code for the World class in the code cell below.

Define the following methods:
* `add` will be passed an object and a location; it should store the object in the specified location 
* `fetch` should return the current contents of a location (which may be `None`)
* `remove` should set a location to `None`.

Finally, define a method named `biota` that will return a list of all items currently in the world.

##### <span style=color:red>Code</span> 

**Important** Write the definition of your World class in the following code cell. Do not delete or move this cell.

<!-- grade_id: world, points: 10, grade: true, solution: true -->
```python
class World:
    """
    A World object represents a grid object that contains a length and width.
    It is given rows and columns with the default value for each element being None.
    """
    
    def __init__(self, row, column):
        self._row = row
        self._column = column
        self._w = np.full([row, column], None)
        
    def __repr__(self):
        return np.array_repr(self._w)
    
    def __getitem__(self, loc):
        return self._w[loc]
    
    def __setitem__(self, loc, val):
        self._w[loc] = val
    
    def add(self, obj, loc):
        """
        Stores the given value in a specified location.
        Parameters:
            self: Mandatory value in order to access values
            obj: The desired object to be placed
            loc: The location where the object is to be placed
        Returns:
            None: sets location equal to an object
        Rtype:
            :None
        """
        self._w[loc] = obj
    
    def fetch(self, loc):
        """
        Returns the object at the position specified.
        Parameters:
            self: Mandatory value in order to access values
            loc: The location where the object is to be retrieved
        Returns:
            :The object at the given location, can be None if nothing is in position
        Rtype:
            :Animal Object
        """
        return self._w[loc]
    
    def remove(self, item, loc):
        """
        Sets the given location to none if the item 
        Parameters:
            self: Mandatory value in order to access values
            item: the item at said location
            loc: The location where the object is to be retrieved
        Returns:
            None: removes the object from the location in the World grid
        Rtype:
            :None
        """
        if self._w[loc] == item:
            self._w[loc] = None

    def biota(self):
        """
        Returns a list of all the objects within the World object Grid.
        Parameters:
            self: Mandatory value in order to access values within class
        Returns:
            li: A list of all the objects in the World object Grid
        Rtype:
            List: A list object
        """
        li = []
        for i in range(self._row):
            for j in range(self._column):
                if self._w[i, j] != None:
                    li.append(self._w[i, j])
        return li

```

##### <span style="color:red">Tests to ensure code is correct:</span>

<!-- grade_id: world_test_1, points: 1, grade: true, solution: false -->
```python
# A new world has no objects
w0 = World(5,5)

assert len(w0.biota()) == 0
```

<!-- grade_id: world_test_2, points: 2, grade: true, solution: false -->
```python
# Test the add and fetch operators
w1 = World(5,5)

w1.add('hello', (0,0))
w1.add('world', (1,1))

assert w1.fetch((0,0)) == 'hello'
assert w1.fetch((1,1)) == 'world'
```

<!-- grade_id: world_test_3, points: 2, grade: true, solution: false -->
```python
# Test the remove and biota methods
w2 = World(5,5)

w2.add('hello', (0,0))
w2.add('world', (1,1))

assert sorted(w2.biota()) == ['hello','world']

w2.remove('hello', (0,0))
w2.remove('world', (1,1))

assert w2.biota() == []
```

<!-- grade_id: world_doc, points: 5, grade: true, solution: true -->
Grader Comments Cell:

## <span style="color:teal;">Fish and Bears</span>

To run a simulation we need to add a random collection of animals to the world.  The two types of animals in this simulation are fish and bears, and you will write class definitions named Fish and Bear that implement the behaviors of the animals.

### Fish Class

During the simulation a Fish object needs to behave as follows:

(1) Fish are susceptible to overcrowding:  if there are fish in 2 or more neighboring cells the fish dies (it's removed from the simulation)

(2) A fish can reproduce if it has been alive for a certain number of time steps: a random neighboring cell is chosen, and if that cell is empty, a new fish is placed in that cell

(3) A fish can move to another cell:  it picks a random direction, and if the neighboring cell in that direction is unoccupied the fish moves there

The class should include the following methods:
* `live` implements rules 1 and 2 shown above
* `move` implements rule 3
* `location` returns the current grid location (row and column) of the object

Define a class variable named `breed_interval` to specify how many time steps a fish must be alive before it reproduces; the initial value for this variable is 12.

### Bear Class

During the simulation a Bear object needs to behave as follows:

(1) A bear looks for fish in each adjacent cell; if it finds one or more fish it eats one at random 

(2) If a bear has not eaten for certain number of time steps it dies (it's removed from the simulation)

(3) A bear can reproduce if it has been alive for a certain number of time steps: a random neighboring cell is chosen, and if that cell is empty, a new bear is placed in that cell

(4) A bear can move to another cell:  it picks a random direction, and if the neighboring cell in that direction is unoccupied the bear moves there

Define a class variable named `breed_interval` to specify how many time steps a bear must be alive before it reproduces; the initial value for this variable is 8.  Define another class variable named `survive_without_food` to be the number of time steps a bear can live before it dies from starvation; the initial value for this variable is 10.

The class should include the following methods:
* `live` implements rules 1, 2, and 3 shown above
* `move` implements rule 4
* `location` returns the current grid location (row and column) of the object

### Animal (Base Class)?

### Details and Hints

**Incremental Development:**  One strategy you might consider is to write a complete implementation for one class, either Bear or Fish.  After you have debugged the class and it passes its unit tests you'll have a better idea of what to implement in the other class.  Then you can start moving common behaviors to the Animal class while you are writing the other derived class.

**Fish:** &nbsp; The way the simulation is defined a fish might be eaten before the top level simulation calls the `live` method.  Your `live` method should check to make sure the fish is still alive.  The easiest way to do this is to include an instance variable named `_alive` that is set to True when the fish is initialized and set to False when it dies.

You can have additional instance methods, class variables, or class methods.  Make sure you describe any new additions in the documentation.

**Animal:** &nbsp; Keep this in mind when you design the class hierarchy: every object has an attribute named ``__class__`` (with two underscores before and after the name).  It is a reference to the class an object was defined with.  As an example of how to use it, consider what would happen if we want to define `reproduce` in the Animal class so it is inherited by both Fish and Bear classes.  We need to know the value of `breed_interval` in each object's own class.  We can find this value using the expression
```
self.__class__.breed_interval
```
This will be a reference to `Fish.breed_interval` or `Bear.breed_interval`, depending on whether a `reproduce` was called with Fish object or Bear object.


### <span style="color:teal;">Animal (15 points)</span>

From the descriptions above it should be apparent that fish and bears have some things in common.

Define a class named Animal and use it as the base class for your Fish and Bear classes.  Some things to think about as you design your classes:
* are there behaviors or operations that are common to both and that can be implemented just once in Animal?
* perhaps a behavior or operation can be defined with a default in the Animal class, and then overridden in the derived class?


If you implement the Animal class write the definition in the following code cell

**Important:** &nbsp; Do not delete or move this cell.

<!-- grade_id: animal, points: 10, grade: true, solution: true -->
```python
class Animal(World):
    """
    An Animal class is used as the base class for the fish and bear objects that will be created later.
    The animal class is within the World object. 
    """
    global_id=0
    
    def __init__(self, world, cords):
        world.add(self,(cords))
        self._w = world
        self._location = cords
        Animal.global_id+=1
        self._name = Animal.global_id
        
    def __getitem__(self, loc):
        return self._w[loc]
    
    def __setitem__(self, loc, val):
        self._w[loc] = val
    
    def location(self):
        """
        Returns location of Animal object
        Parameters:
            self: mandatory value in order to access values
        Returns:
            :location of the Animal object
        Rtype:
            :None
        """
        return self._location
    
    def random_walk(self):
        """
        Provides a random
        Parameters:
            self: mandatory value in order to access values
        Returns:
            :location of the Animal object
        Rtype:
            
        """
        return self.get_move_dict()[randint(1,8)]
    
    def move(self):
        """
        Returns location of Animal object after it is moved to a random location
        Parameters:
            self: mandatory value in order to access values
        Returns:
            :location of the Animal object after being randomly moved
        Rtype:
        """
        new_location=self.random_walk()
        while self._w.fetch(new_location) != None:
            new_location=self.random_walk()
        self._w.remove(self,self._location)
        self._w.add(self,new_location)
        self._location=new_location
        return self._location
    
    def get_move_dict(self):
        """
        Creates a dictionary of all possible locations an Animal object can move in one step within the grid.
        Parameters:
            self: mandatory value in order to access values
        Returns:
            self.move_dict: dictionary of all possible locations for the object to move
        Rtype: 
            Dictionary
        """
        rows = self._w._row
        colums = self._w._column
        x,y = self._location
        self.move_dict = {1:(x,y-1),
                     2:(x,(y+1)%(rows)),
                     3:(x-1,y),
                     4:((x+1)%(colums),y),
                     5:(x-1,y-1),
                     6:(x-1,(y+1)%(rows)),
                     7:((x+1)%(colums),y-1),
                     8:((x+1)%(colums),(y+1)%(rows))
                     }
        return self.move_dict
```

<!-- grade_id: animal_doc, points: 5, grade: true, solution: true -->
Grader Comments Cell:

### <span style="color:teal;">Fish (20 points)</span>

<!-- grade_id: fish, points: 7, grade: true, solution: true -->
```python
class Fish(Animal):
    """
    The Fish class is a subclass of the Animal class. It is an object found within the
    World grid. 
    """
    breed_interval=12
    
    def __init__(self, world, coords):
        Animal.__init__(self, world, coords)
        self._alive = True
        
    def __repr__(self):
        return "{}{}".format(u'\U0001f41f',self._name)
    
    def _check_around(self,location):
        """
        Checks around the Fish object at the given location and reteurns empty spots
        Parameters:
            self: mandatory value in order to access value
            location: loation of Fish object
        Return: 
            count: returns count of empty spaces around Fish object
        Rtype: 
            integer
        """
        x,y=location
        count=0
        rows = self._w._row
        colums = self._w._column
        move_dict = Animal.get_move_dict(self)
        for key,value in move_dict.items():
            if self._w.fetch(value)!=None:
                count += 1
        return count
        
    def live(self):
        """
        Determines whether the Fish object is alive or not.
        Parameters:
            self: mandatory value in order to access values
        Return:
            None: sets whether the fish is alive or dead in boolean form
        Rtype:
            None
        """
        self.breed_interval-=1
        if self._alive ==True:
            if self._check_around(self._location)>=2:
                self._w.remove(self,self._location)
                self._alive=False
            if self.breed_interval==0:
                breed_location=self.__class__.random_walk(self)
                while self._w.fetch(breed_location) != None:
                    breed_location=self.__class__.random_walk(self)
                Fish(self._w, breed_location)
                self.breed_interval=12
```

##### <span style="color:red">More Test:</span>

<!-- grade_id: fish_test_1, points: 2, grade: true, solution: false -->
```python
# A new world has no objects
fw1 = World(5,5)
assert len(fw1.biota()) == 0

# After adding a fish there should be one object
f1 = Fish(fw1, (2,2))
assert len(fw1.biota()) == 1

# Test the location method
assert f1.location() == (2,2)
```

<!-- grade_id: fish_test_2, points: 2, grade: true, solution: false -->
```python
# Setting breed_interval to 1 should cause a fish to reproduce when live is called
fw2 = World(5,5)
Fish.breed_interval = 1
f2 = Fish(fw2, (2,2))
f2.live()
assert len(fw2.biota()) == 2

# Reset the interval to original value for remaining tests
Fish.breed_interval = 12
```

<!-- grade_id: fish_test_3, points: 2, grade: true, solution: false -->
```python
# Make three fish, the one in the middle should die from overcrowding
fw3 = World(5,5)
f3 = Fish(fw3, (2,2))
Fish(fw3, (1,1))
Fish(fw3, (3,3))
f3.live()
assert len(fw3.biota()) == 2
```

<!-- grade_id: fish_test_4, points: 2, grade: true, solution: false -->
```python
# When a fish moves it should be within one cell of its original location
fw4 = World(5,5)
f4 = Fish(fw4, (2,2))
f4.move()
r, c = f4.location()
assert (r,c) != (2,2)
assert abs(r-2) <= 1 and abs(c-2) <= 1
```

<!-- grade_id: fish_doc, points: 5, grade: true, solution: true -->
Grader Comments Cell:

###  <span style="color:teal;">Bear (20 points)</span>

<!-- grade_id: bear, points: 7, grade: true, solution: true -->
```python
class Bear(Animal):
    """
    The Bear class is a subclass of the Animal class. It is an object found within the
    World grid.
    """
    breed_interval=8
    survive_without_food=10
    
    def __init__(self, world, coords):
        Animal.__init__(self, world, coords)
        self._alive=True

    def __repr__(self):
        return "{}{}".format(u'\U0001f43b',self._name)
    
    def _nearby_food(self,location):
        """
        Searches blocks adjacent to the specified Bear objects position for a Fish object as food.
        Parameters:
            self: mandatory value in order to access values
            location: location of the Bear object
        Returns:
            options_list: list of possible options for which fish the Bear object should eat
        Rtype:
            List object
        """
        x,y=location
        options_list=[]
        rows = self._w._row
        colums = self._w._column
        move_dict = Animal.get_move_dict(self)
        for key,value in move_dict.items():
            if type(self._w.fetch(value))==Fish:
                options_list.append(value)
        return options_list
                
    def _eat(self):
        """
        Determines whether there are multiple possible food options and randomly selects which fish to eat.
        Parameters:
            self: mandatory value in order to access values
        Returns: 
            None: eats a Fish object and removes it from the World grid
        Rtype:
            :None
        """
        food_options = self._nearby_food(self._location)
        if len(food_options)>0:
            chosen_fish_loc=food_options[randint(0,len(food_options)-1)]
            chosen_fish=self._w.fetch(chosen_fish_loc)
            chosen_fish._alive=False
            self._w.remove(chosen_fish,chosen_fish_loc)
            self.survive_without_food=10
    
    def live(self):
        """
        Determines whether the Bear object is alive or not.
        Parameters:
            self: mandatory value in order to access values
        Return:
            None: sets whether the fish is alive or dead in boolean form
        Rtype:
            :None
        """
        self.breed_interval-=1
        self.survive_without_food-=1
        if self.survive_without_food==0:
            self._alive=False
            self._w.remove(self,self._location)
        if self._alive ==True:
            self._eat()
            if self.breed_interval==0:
                breed_location=self.__class__.random_walk(self)
                while self._w.fetch(breed_location) != None:
                    breed_location=self.__class__.random_walk(self)
                Bear(self._w, (breed_location))
                self.breed_interval=8
```

##### <span style="color:red">More Tests:</span>

<!-- grade_id: bear_test_1, points: 2, grade: true, solution: false -->
```python
# Test the Bear constructor and location method
bw1 = World(5,5)
b1 = Bear(bw1, (1,1))
assert len(bw1.biota()) == 1
assert b1.location() == (1,1)
```

<!-- grade_id: bear_test_2, points: 2, grade: true, solution: false -->
```python
# Repeat the reproduction test for Bears
bw2 = World(5,5)
Bear.breed_interval = 1
b2 = Bear(bw2, (2,2))
b2.live()
assert len(bw2.biota()) == 2
Bear.breed_interval = 8
```

<!-- grade_id: bear_test_3, points: 2, grade: true, solution: false -->
```python
# Make fish for the bear to eat, count the number of objects after eating
bw3 = World(5,5)
b3 = Bear(bw3, (2,2))
Fish(bw3, (1,1))
Fish(bw3, (3,3))
b3.live()
assert len(bw3.biota()) == 2
```

<!-- grade_id: bear_test_4, points: 2, grade: true, solution: false -->
```python
# Setting the survival limit to 1 should cause a Bear to starve 
bw4 = World(5,5)
Bear.survive_without_food = 1
b4 = Bear(bw4, (2,2))
b4.live()
assert len(bw4.biota()) == 0
Bear.survive_without_food = 10
```

<!-- grade_id: bear_test_5, points: 2, grade: true, solution: false -->
```python
# Repeat the move test for bears
bw5 = World(5,5)
b5 = Bear(bw5, (2,2))
b5.move()
r, c = b5.location()
assert (r,c) != (2,2)
assert abs(r-2) <= 1 and abs(c-2) <= 1
```

<!-- grade_id: bear_doc, points: 3, grade: true, solution: true -->
Grader Comments Cell:

##  <span style="color:teal;">The `wbf` Function</span>

The `wbf` function so it returns a new World object with the specified number of rows and columns and with the specified number of Bear and Fish objects at random locations.

When we grade your project we will call `wbf` to make a World object and then use the main loop (implemented by `step_system`) to run the simulation.

<!-- grade_id: wbf, points: 9, grade: true, solution: true -->
```python
def wbf(nrows, ncols, nbears, nfish):
    """
    Returns a World object with the specified number of Bears and Fish objects randomly throughout the grid with specified
    length and width.
    Paramters:
        nrows: number of rows in the new World object
        ncols: number of columns in the new World object
        nbears: number of Bear objects in the new World object
        nfish: number of Fish objects in the new World object
    Returns:
        world: A new World object with the specified number of Bears and Fish objects randomly throughout the grid
    rtype:
        World Object
    """
    bear_count=0
    fish_count=0
    world= World(nrows,ncols)
    
    while bear_count<nbears:
        bear_cord = (randint(0,nrows-1),randint(0,ncols-1))
        if world.fetch(bear_cord)==None:
            Bear(world,(bear_cord))
            bear_count += 1
            
    while fish_count<nfish:
        fish_cord = (randint(0,nrows-1),randint(0,ncols-1))
        if world.fetch(fish_cord)==None:
            Fish(world,(fish_cord))
            fish_count += 1
        
    return world
```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: wbf_test, points: 3, grade: true, solution: false -->
```python
w = wbf(10,10,3,12)
dct = { Bear: 0, Fish: 0 }
for x in w.biota():
    dct[x.__class__] += 1

assert dct[Bear] == 3
assert dct[Fish] == 12
```

<!-- grade_id: wbf_doc, points: 3, grade: true, solution: true -->
Grader Comments Cell:

##  <span style="color:teal;">The `step_system` Function</span>

```python
def step_system(world):
    for x in world.biota():
        x.live()
    for x in world.biota():
        x.move()
```

```python
if IPython:
    logging = True
    w = wbf(10,10,3,12)
    for i in range(3):
        print(w)
        step_system(w)
```

Output:
```
array([[None, None, None, None, None, None, None, None, 🐟69, None],
       [None, None, None, None, 🐟67, None, None, None, None, None],
       [None, None, None, None, None, None, None, None, None, None],
       [None, None, None, None, None, None, None, None, None, None],
       [None, None, None, None, None, None, None, None, 🐟72, None],
       [None, None, 🐟66, None, None, None, None, 🐟68, None, None],
       [🐻63, None, None, None, 🐟71, None, None, None, None, 🐟64],
       [🐟65, None, None, None, None, 🐟74, None, None, None, None],
       [None, None, None, 🐟75, 🐻61, 🐻62, 🐟70, None, None, None],
       [None, None, 🐟73, None, None, None, None, None, None, None]], dtype=object)
array([[None, 🐟73, None, None, None, None, None, None, None, None],
       [None, None, None, None, None, None, None, None, 🐟69, None],
       [None, None, None, None, None, 🐟67, None, None, None, None],
       [None, None, None, None, None, None, None, None, None, None],
       [None, None, None, None, None,
```

<!-- grade_id: experiments, points: 10, grade: true, solution: true -->
Grader Comments Cell: